## Before submitting
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

2. Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.

3. Please submit only the `*.ipynb` file.

4. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Edit only between `YOUR CODE HERE` and `END YOUR CODE`.

5. Make sure to use Python 3, not Python 2.

Fill your group name and collaborators below:

```python
GROUPNAME = ""
COLLABORATORS = ""
```

---

# Exercise Sheet 1: Python Basics

This first  exercise sheet tests the basic functionalities of the Python programming language in the context of a simple prediction task. We consider the problem of predicting health risk of subjects from personal data and habits. We first use for this task a decision tree

![](tree.png)

adapted from the webpage http://www.refactorthis.net/post/2013/04/10/Machine-Learning-tutorial-How-to-create-a-decision-tree-in-RapidMiner-using-the-Titanic-passenger-data-set.aspx. For this exercise sheet, you are required to use only pure Python, and to not import any module, including numpy. In exercise sheet 2, the nearest neighbor part of this exercise sheet will be revisited with numpy.

## Classifying a single instance (15 P)

* Create a function that takes as input a tuple containing values for attributes (smoker,age,diet), and computes the output of the decision tree. Should return `"less"` or `"more"`.
* Test your function on the tuple `('yes', 31, 'good')`,

<!-- grade_id: cell-b66d7278bc313c94, points: 0, grade: false, solution: true -->
```python
def decision(x):
    isSmoker = x[0]
    age = x[1]
    isDietGood = x[2]
    
    # >>>>> YOUR CODE HERE
    if isSmoker == 'yes':
        if age < 29.5:
            return 'less'
        else: 
            return 'more'
    elif isDietGood == 'good':
        return 'less'
    else: 
        return 'more'
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-c31b80471db3132f, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert decision(x) == 'more'

```

```python
1 and 0 
```

Output:
```
0
```

## Reading a dataset from a text file (10 P)

The file `health-test.txt` contains several fictious records of personal data and habits.

* Read the file automatically using the methods introduced during the lecture.
* Represent the dataset as a list of tuples. Make sure that the tuples have the same format as above, e.g. `('yes', 31, 'good')`.
* Make extra note of the datatype of each element

<!-- grade_id: cell-c1a8bc4c0e4ccb26, points: 0, grade: false, solution: true -->
```python
import numpy as np
import numpy.random as rng
import matplotlib.pyplot as plt
```

```python
gettest()
```

Output:
```
[('yes', 21, 'poor'),
 ('no', 50, 'good'),
 ('no', 23, 'good'),
 ('yes', 45, 'poor'),
 ('yes', 51, 'good'),
 ('no', 60, 'good'),
 ('no', 15, 'poor'),
 ('no', 18, 'good')]
```

<!-- grade_id: cell-4e1f7ad1e66b3121, points: 10, grade: true, solution: false -->
```python

```

## Applying the decision tree to the dataset (15 P)

* Apply the decision tree to all points in the dataset, and return the ratio of them that are classified as "more".

<!-- grade_id: cell-6703ef98e2b5c93b, points: 0, grade: false, solution: true -->
```python
def evaluate_testset():
    more_counter = 0
    data_set = gettest()
    for data in data_set:
        if decision(data) == 'more':
            more_counter += 1
    ratio = more_counter / len(data_set)
    return ratio
```

```python
evaluate_testset()
```

Output:
```
0.375
```

<!-- grade_id: cell-c13a0b23c9faba52, points: 15, grade: true, solution: false -->
```python

```

## Learning from examples (10 P)

Suppose that instead of relying on a fixed decision tree, we would like to use a data-driven approach where data points are classified based on a set of training observations manually labeled by experts. Such labeled dataset is available in the file `health-train.txt`. The first three columns have the same meaning than for `health-test.txt`, and the last column corresponds to the labels.

* Write a procedure that reads this file and converts it into a list of pairs. The first element of each pair is a triplet of attributes, and the second element is the label.

<!-- grade_id: cell-fc38ed11fee6fbeb, points: 0, grade: false, solution: true -->
```python
def gettrain():
    with open('health-train.txt', 'r') as file:
        data = list()
        for line in file:
            elements = str.split(line[:-1], ',')
            
            data_tupel = (elements[0], int(elements[1]), elements[2])
            data.append((data_tupel, elements[3]))
        return data
        
```

```python
gettrain()
```

Output:
```
[(('yes', 54, 'good'), 'less'),
 (('no', 55, 'good'), 'less'),
 (('no', 26, 'good'), 'less'),
 (('yes', 40, 'good'), 'more'),
 (('yes', 25, 'poor'), 'less'),
 (('no', 13, 'poor'), 'more'),
 (('no', 15, 'good'), 'less'),
 (('no', 50, 'poor'), 'more'),
 (('yes', 33, 'good'), 'more'),
 (('no', 35, 'good'), 'less'),
 (('no', 41, 'good'), 'less'),
 (('yes', 30, 'poor'), 'more'),
 (('no', 39, 'poor'), 'more'),
 (('no', 20, 'good'), 'less'),
 (('yes', 18, 'poor'), 'less'),
 (('yes', 55, 'good'), 'more')]
```

<!-- grade_id: cell-a3d593f232e0403a, points: 10, grade: true, solution: false -->
```python

```

## Nearest neighbor classifier (25 P)

We consider the nearest neighbor algorithm that classifies test points following the label of the nearest neighbor in the training data. For this, we need to define a distance function between data points. We define it to be

`d(a, b) = (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])`

where `a` and `b` are two tuples corrsponding to the attributes of two data points.

* Write a function that retrieves for a test point the nearest neighbor in the training set, and classifies the test point accordingly.
* Test your function on the tuple `('yes', 31, 'good')`

<!-- grade_id: cell-671ea24ec8a11241, points: 0, grade: false, solution: true -->
```python
def neighbor(x, trainset):
    error_list = list()
    for t in trainset:
        error = (x[0] != t[0][0]) + ((x[1] - t[0][1]) / 50.0) ** 2 + (x[2] != t[0][2])
        error_list.append(error)
    index = 0
    for i, v in enumerate(error_list):
        if min(error_list) == v:
            index = i
    print(trainset[index][0])
    return trainset[index][1]

```

```python
neighbor(('yes', 31, 'good'), gettrain())
```

Output:
```
('yes', 33, 'good')
'more'
```

<!-- grade_id: cell-a36122337853f195, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert neighbor(x, gettrain()) == "more"

```

* Apply both the decision tree and nearest neighbor classifiers on the test set, and return the list of data point(s) for which the two classifiers disagree, and with which probability it happens.
* A data point should look like above, e.g. `('yes', 31, 'good')`.

<!-- grade_id: cell-8dbf7da153f3d797, points: 0, grade: false, solution: true -->
```python
def compare():
    testset = gettest()
    trainset = gettrain()
    Xdisagree = list()
    disagree_counter = 0
    for data in testset:
        neighbor_result = neighbor(data, trainset)
        decision_result = decision(data)
        if neighbor_result != decision_result:
            disagree_counter += 1
            Xdisagree.append(data)
    probability = disagree_counter / len(testset)
    
    return Xdisagree, probability
```

```python
compare()
```

Output:
```
([('yes', 51, 'good')], 0.125)
```

<!-- grade_id: cell-3b55f7e89ad4dfeb, points: 10, grade: true, solution: false -->
```python
Xdisagree, probability = compare()
assert type(Xdisagree) == list
print(Xdisagree)
assert probability >= 0.0 and probability <= 1.0
```

Output:
```
[('yes', 51, 'good')]
```

One problem of simple nearest neighbors is that one needs to compare the point to predict to all data points in the training set. This can be slow for datasets of thousands of points or more. Alternatively, some classifiers train a model first, and then use it to classify the data.

## Nearest mean classifier (25 P)

We consider one such trainable model, which operates in two steps:

(1) Compute the average point for each class, (2) classify new points to be of the class whose average point is nearest to the point to predict.

For this classifier, we convert the attributes smoker and diet to real values (for smoker: yes=1.0 and no=0.0, and for diet: good=0.0 and poor=1.0), and use the modified distance function:

`d(a,b) = (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2`

We adopt an object-oriented approach for building this classifier.

* Implement the methods `train` and `predict` of the class `NearestMeanClassifier`.

<!-- grade_id: cell-e0b339bfd0fcc16c, points: 0, grade: false, solution: true -->
```python
class NearestMeanClassifier:
    def train(self, dataset):
        # classes
        more = list()
        less = list()

        # transform the isSmoker and isDietGood in number
        for i, data in enumerate(dataset):
            smoker = 1.0 if data[0][0] == 'yes' else 0.0 
            diet = 1.0 if data[0][2] == 'good' else 0.0
            point = (smoker, data[0][1], diet)
            if data[1] == 'more':
                more.append(point)
            else: 
                less.append(point)
                
        lessAvgSmoker = 0
        lessAvgAge = 0
        lessAvgDiet = 0
        for point in less:
            lessAvgSmoker += point[0]
            lessAvgAge += point[1]
            lessAvgDiet += point[2]
            
        self.lessAvg = (lessAvgSmoker / len(less), lessAvgAge / len(less), lessAvgDiet / len(less))

        moreAvgSmoker = 0
        moreAvgAge = 0
        moreAvgDiet = 0
        for point in more:
            moreAvgSmoker += point[0]
            moreAvgAge += point[1]
            moreAvgDiet += point[2]
            
        self.moreAvg = (moreAvgSmoker / len(more), moreAvgAge / len(more), moreAvgDiet / len(more))

        # probably one point for more or value ?? 
        # probably one point for less 
        print(self.moreAvg)
        print(self.lessAvg)
        

    def predict(self, x):
        smoker = 1.0 if x[0] == 'yes' else 0.0 
        diet = 1.0 if x[2] == 'good' else 0.0
        point = (smoker, x[1], diet)
        # if x is closer to average of more or less 
        errorMore = (point[0] - self.moreAvg[0]) ** 2 + 
        ((point[1] - self.moreAvg[1]) / 50.0) ** 2 + (point[2] - self.moreAvg[2]) ** 2
        errorLess = (point[0] - self.lessAvg[0]) ** 2 + 
        ((point[1] - self.lessAvg[1]) / 50.0) ** 2 + (point[2] - self.lessAvg[2]) ** 2
        
        prediction = ''
        if errorMore < errorLess:
            prediction = 'more'
        else: 
            prediction = 'less'
        # prediction = 'more' if x > self.avg else 'below'
        # if less or more
        return prediction
```

* Build an object of class `NearestMeanClassifier`, train it on the training data, and return the mean vector for each class. You should return a dictionary with keys `less` and `more`. Each key should correspond to a mean vector.

<!-- grade_id: cell-5f7f00ee83c94703, points: 0, grade: false, solution: true -->
```python
def build_and_train():
    train_data = gettrain()
    test_data = gettest()
    
    # train
    classifier = NearestMeanClassifier()
    classifier.train(train_data)
    # predict
    classification = ''
    more = 0
    less = 0
    for i, data in enumerate(test_data):
        classification = classifier.predict(data)
        print(classification)
        if classification == 'more':
            more += 1
        else: 
            less += 1
    
    mean = {
        'less': less/len(test_data),
        'more': more/len(test_data)
    }
build_and_train()
```

Output:
```
(0.5714285714285714, 37.142857142857146, 0.42857142857142855)
(0.3333333333333333, 32.111111111111114, 0.7777777777777778)
more
less
less
more
more
less
more
less
```

<!-- grade_id: cell-415891bde4cbde19, points: 15, grade: true, solution: false -->
```python

```

* Predict the test data using the nearest mean classifier and return all test examples for which all three classifiers (decision tree, nearest neighbor and nearest mean) agree. You should return a list containing tuples. Each tuple should contain the datapoint and the prediction i.e.

`[(('no', 50, 'good'), 'less'), ... ]`

<!-- grade_id: cell-f37f3035a32a8f85, points: 0, grade: false, solution: true -->
```python
def predict_test():
    agreed_samples = list()
    test_data = gettest()
    train_data = gettrain()
    print(test_data)
    print(train_data)
    classifier = NearestMeanClassifier()
    classifier.train(train_data)
    
    for i, data in enumerate(test_data):
        neighbor_result = neighbor(data, train_data)
        decision_result = decision(data)
        classifier_result = classifier.predict(data)
        if decision_result == decision_result and decision_result == classifier_result:
            agreed_samples.append((data, neighbor_result))
    return agreed_samples
                                  
predict_test()
```

Output:
```
[('yes', 21, 'poor'), ('no', 50, 'good'), ('no', 23, 'good'), ('yes', 45, 'poor'), ('yes', 51, 'good'), ('no', 60, 'good'), ('no', 15, 'poor'), ('no', 18, 'good')]
[(('yes', 54, 'good'), 'less'), (('no', 55, 'good'), 'less'), (('no', 26, 'good'), 'less'), (('yes', 40, 'good'), 'more'), (('yes', 25, 'poor'), 'less'), (('no', 13, 'poor'), 'more'), (('no', 15, 'good'), 'less'), (('no', 50, 'poor'), 'more'), (('yes', 33, 'good'), 'more'), (('no', 35, 'good'), 'less'), (('no', 41, 'good'), 'less'), (('yes', 30, 'poor'), 'more'), (('no', 39, 'poor'), 'more'), (('no', 20, 'good'), 'less'), (('yes', 18, 'poor'), 'less'), (('yes', 55, 'good'), 'more')]
(0.5714285714285714, 37.142857142857146, 0.42857142857142855)
(0.3333333333333333, 32.111111111111114, 0.7777777777777778)
('yes', 18, 'poor')
('no', 55, 'good')
('no', 20, 'good')
('yes', 30, 'poor')
('yes', 54, 'good')
('no', 55, 'good')
('no', 13, 'poor')
('no', 20, 'good')
[(('no', 50, 'good'), 'less'),
 (('no', 23, 'good'), 'less'),
 (('yes'
```

<!-- grade_id: cell-853c957eaaf81c28, points: 10, grade: true, solution: false -->
```python

```

```python

```